# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu: IAD25
Członkowie: Karolina Minor, Filip Przyczyna, Maciej Zabdyr, Maria Zalewska

---

## Zestaw zadań 3: Elementy statystyki opisowej 3 (Notatnik 1/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Przygotowanie danych [10 pkt]

#### a) Natężenie ruchu rowerowego 

Plik raw_bike_traffic_data.csv zawiera dane na temat ruchu rowerowego na moście Fremont Bridge w Seattle pobrane ze strony
https://data.seattle.gov/Transportation/Fremont-Bridge-Bicycle-Counter/65db-xm6k.

Zaimportuj zawarte w pliku dane do notatnika i utwórz na ich podstawie tablicę `bike_traffic_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > `pandas.PeriodIndex` zawierający okresy czasu, w których dokonywane były zliczenia rowerzystów.
    
    > Kolumny: Total, East, West</br>
    > Kolumny zawierają liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.
- Pozostaw w tablicy wyłącznie wiersze dotyczące 2022 roku.

Zapisz tablicę do pliku bike_traffic_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 12 pierwszych i ostatnich wierszy, a zawarte w tablicy wartości wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-aa85199b5cb633d2, points: 0, grade: false, solution: true -->
```python
bike_traffic_data = pd.read_csv( 'raw_data/raw_bike_traffic_data.csv', index_col = 'Date' )
bike_traffic_data = bike_traffic_data.rename( columns = { 'Fremont Bridge Total': 'Total', 'Fremont Bridge East Sidewalk': 'East', 'Fremont Bridge West Sidewalk': 'West' } )
bike_traffic_data.index = pd.to_datetime( bike_traffic_data.index )
bike_traffic_data.index = bike_traffic_data.index.to_period( 'H' )
bike_traffic_data = bike_traffic_data[ bike_traffic_data.index.year == 2022 ]
bike_traffic_data.to_pickle( 'bike_traffic_data.pkl' )

combined_rows = pd.concat( [ bike_traffic_data.head( 12 ), bike_traffic_data.tail( 12 ) ] )
print( combined_rows )
```

Output:
```
/tmp/ipykernel_8864/665272851.py:3: UserWarning: Could not infer format, so each element will be parsed individually, falling back to `dateutil`. To ensure parsing is consistent and as-expected, please specify a format.
  bike_traffic_data.index = pd.to_datetime( bike_traffic_data.index )
                  Total  East  West
Date                               
2022-01-01 00:00    0.0   0.0   0.0
2022-01-01 01:00    0.0   0.0   0.0
2022-01-01 02:00    1.0   0.0   1.0
2022-01-01 03:00    0.0   0.0   0.0
2022-01-01 04:00    0.0   0.0   0.0
2022-01-01 05:00    1.0   1.0   0.0
2022-01-01 06:00    1.0   1.0   0.0
2022-01-01 07:00    1.0   1.0   0.0
2022-01-01 08:00    3.0   3.0   0.0
2022-01-01 09:00    2.0   0.0   2.0
2022-01-01 10:00    0.0   0.0   0.0
2022-01-01 11:00    1.0   0.0   1.0
2022-12-31 12:00  124.0  63.0  61.0
2022-12-31 13:00  103.0  56.0  47.0
2022-12-31 14:00   98.0  46.0  52.0
2022-12-31 15:00   83.0  40.0  43.0
2022-12-31 16:00   68.0  28.0  40.0
2022-12-31 17:00   35.0  1
```

<!-- grade_id: cell-683b0a3c6c74d0c6, points: 2, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert bike_traffic_data.shape == (8760, 3)
# Test 2 Kolumny
assert len(set(list(bike_traffic_data.columns)) & set(['Total', 'East', 'West'])) == 3
# Test 3 Indeks
assert (bike_traffic_data.index.name == 'Date') == True
assert isinstance(bike_traffic_data.index, pd.PeriodIndex)
assert bike_traffic_data.index.values.min() == pd.Period('2022-01-01 00:00', 'H')
assert bike_traffic_data.index.values.max() == pd.Period('2022-12-31 23:00', 'H')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].min().values, np.array([0, 0, 0]))) == True
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].max().values, np.array([717, 464, 635]))) == True
```

Output:
```
/tmp/ipykernel_8864/1841253091.py:9: FutureWarning: 'H' is deprecated and will be removed in a future version, please use 'h' instead.
  assert bike_traffic_data.index.values.min() == pd.Period('2022-01-01 00:00', 'H')
/tmp/ipykernel_8864/1841253091.py:10: FutureWarning: 'H' is deprecated and will be removed in a future version, please use 'h' instead.
  assert bike_traffic_data.index.values.max() == pd.Period('2022-12-31 23:00', 'H')
```

#### b) Pogoda

Pliki 2022_01.txt, 2022_02.txt ...  2022_12.txt zawierają miesięczne raporty pogodowe (F6 Preliminary Local Climate Data Report) ze stacji NWS Seattle/Tacoma, WA pobrane ze strony https://forecast.weather.gov/product.php?site=SEW&product=CF6&issuedby=SEW.

Wyjaśnienie wykorzystanych w raportach symboli i skrótów można znaleźć pod następującym linkiem: https://www.weather.gov/grr/climateF6explain.

Zaimportuj zawarte w plikach dane do notatnika i utwórz na ich podstawie tablicę `weather_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > `pandas.PeriodIndex` zawierający okresy czasu, których dotyczą dane z raportów pogodowych.</br>
    
    > Kolumny: 2-poziomowe nagłówki w podanej niżej kolejności (nazwa 1 poziomu: podległe jej nazwy 2 poziomu) </br>
    >> Temperature: Min, Avg, Max</br>
    >> Precipitation: Total</br>
    >> Snow: Fall, Depth</br>
    >> Wind:  Avg speed, Max speed</br>
    >> Sky: Cloud cover</br>

- Zamień oznaczenia opadów śladowych na wartość 0.
- Przelicz występujące w tablicy wartości do *C (temperatura), mm (opady), m/s (wiatr) i ułamków (zachmurzenie).

Zapisz tablicę do pliku weather_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 10 pierwszych i ostatnich wierszy, a zawarte w tablicy liczby wyświetlały się z dokładnością 2 miejsc po przecinku.

<!-- grade_id: cell-186c2f35c4f9aca6, points: 0, grade: false, solution: true -->
```python
weather_data = []
columns = [ 'DY', 'Max', 'Min', 'Avg', 'DEP', 'HDD', 'CDD', 'Total', 'Fall', 
           'Depth', 'Avg speed', 'Max speed', 'DIR', 'MIN', 'PSBL', 'Cloud cover',
           'WX', 'SPD', 'DR' ]
index_ranges = [
    ( 0, 9 ),
    ( 31, 40 ),
    ( 59, 68 ),
    ( 90, 99 ),
    ( 120, 129 ),
    ( 151, 160 ),
    ( 181, 190 ),
    ( 212, 221 ),
    ( 243, 252 ),
    ( 273, 282 ),
    ( 304, 313 ),
    ( 334, 343 )    
]

k=0
stop = 0
for nr in range(1,13):
    msc = str(nr).zfill(2)
    name = f"2022_{msc}.txt"
    with open(name, "r") as file:
        lines = file.readlines()
        k = 0  # Inicjalizacja zmiennej k
        for ix, line in enumerate(lines):
            if line.strip() == "================================================================================":
                k += 1
            if k == 3:
                stop = ix
                break
    weather_data.extend(lines[18:stop])


weather_data = pd.DataFrame(weather_data)
weather_data[0] = weather_data[0].apply(lambda x: x.replace('T', '0'))
print( weather_data )
for start, end in index_ranges:
    weather_data[0].iloc[start:end] = weather_data[0].iloc[start:end].str.slice(1)
weather_data = weather_data[0].str.split(r'\s+', expand=True)
weather_data = weather_data.iloc[:, :-1]
weather_data.columns = columns


weather_data = weather_data.set_index( 'DY' )
#print( weather_data )

weather_data.index = pd.period_range( start = '2022-01-01', end = '2022-12-31', freq = 'D' )
weather_data.index.name = 'Date'

weather_data['Max'] = pd.to_numeric(weather_data['Max'], errors='coerce')
weather_data['Min'] = pd.to_numeric(weather_data['Min'], errors='coerce')
weather_data['Avg'] = pd.to_numeric(weather_data['Avg'], errors='coerce')
weather_data['Fall'] = pd.to_numeric(weather_data['Fall'], errors='coerce')
weather_data['Depth'] = pd.to_numeric(weather_data['Depth'], errors='coerce')
weather_data['Total'] = pd.to_numeric(weather_data['Total'], errors='coerce')
weather_data['Avg speed'] = pd.to_numeric(weather_data['Avg speed'], errors='coerce')
weather_data['Max speed'] = pd.to_numeric(weather_data['Max speed'], errors='coerce')
weather_data['Cloud cover'] = pd.to_numeric(weather_data['Cloud cover'], errors='coerce')

weather_data['Total'] = weather_data['Total'].replace('T', 0.0)
weather_data['Max'] = ( weather_data['Max'] - 32 ) * 5/9
weather_data['Min'] = ( weather_data['Min'] - 32 ) * 5/9
weather_data['Avg'] = ( weather_data['Avg'] - 32)  * 5/9
weather_data['Total'] *= 25.4
weather_data['Depth'] *= 25.4
weather_data['Fall'] *= 25.4
weather_data['Avg speed'] *= 0.447038889
weather_data['Max speed'] *= 0.447038889
weather_data['Cloud cover'] /= 10

weather_data = weather_data[[ 'Min', 'Avg', 'Max', 'Total', 'Fall', 'Depth', 'Avg speed', 'Max speed', 'Cloud cover' ]]

header = pd.MultiIndex.from_tuples([
    ('Temperature', 'Min'),
    ('Temperature', 'Avg'),
    ('Temperature', 'Max'),
    ('Precipitation', 'Total'),
    ('Snow', 'Fall'),
    ('Snow', 'Depth'),
    ('Wind', 'Avg speed'),
    ('Wind', 'Max speed'),
    ('Sky', 'Cloud cover')
])

weather_data = pd.DataFrame( weather_data.values, index = weather_data.index, columns = header )
weather_data.to_pickle( "weather_data.pkl" )

with pd.option_context( 'display.float_format', '{:0.2f}'.format ):
    print( pd.concat( [ weather_data.iloc[:10], weather_data.iloc[-10:] ] ) )
```

Output:
```
0
0     1  34  22  28 -14  37   0 0.00  0.0    3  8.0...
1     2  43  30  37  -5  28   0 0.79  0.0    3  8.5...
2     3  41  33  37  -5  28   0 0.49    0    0 11.5...
3     4  42  35  39  -3  26   0 0.22  0.0    0  9.7...
4     5  39  34  37  -5  28   0 0.33    0    0  5.0...
..                                                 ...
360  27  53  44  49   7  16   0 0.27  0.0    0 17.3...
361  28  46  41  44   2  21   0 0.03  0.0    0 12.9...
362  29  46  41  44   2  21   0 0.35  0.0    0  6.2...
363  30  53  42  48   6  17   0 0.36  0.0    0 14.4...
364  31  49  43  46   4  19   0 0.01  0.0    0 10.0...

[365 rows x 1 columns]
           Temperature             Precipitation Snow            Wind  \
                   Min   Avg   Max         Total Fall Depth Avg speed   
Date                                                                    
2022-01-01       -5.56 -2.22  1.11          0.00 0.00 76.20      3.58   
2022-01-02       -1.11  2.78  6.11         20.07 0.00 76.20      3.80   
2022
```

<!-- grade_id: cell-e70d426962ca7e21, points: 8, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert weather_data.shape == (365, 9)
# Test 2 Kolumny
assert len(set(list(weather_data.columns)) & set([('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'), ('Precipitation', 'Total'),
            ('Snow', 'Fall'), ('Snow', 'Depth'), ('Wind', 'Avg speed'), ('Wind', 'Max speed'), ('Sky', 'Cloud cover')])) == 9
# Test 3 Indeks
assert (weather_data.index.name == 'Date') == True
assert isinstance(weather_data.index, pd.PeriodIndex) == True
assert weather_data.index.values.min() == pd.Period('2022-01-01', 'D')
assert weather_data.index.values.max() == pd.Period('2022-12-31', 'D')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(weather_data.min().values, np.array([-7.77777778, -5.55555556, -3.33333333, 0., 0., 0., 0.89407778, 3.12927222, 0.]))) == True
assert np.all(np.isclose(weather_data.max().values, np.array([20., 26.66666667, 35., 75.438, 66.04, 76.2, 8.31492333, 16.54043889, 1.]))) == True
```